# Heaps

## Overview

For this assignment you will update the heap data structure implemented in class so that it accepts a `key` function in its initializer, which will allow the contents of the heap to be maintained using an arbitrary priority (as dictated by the key extracted by said function).

See the naive implementation of a priority queue in the [lecture notes](https://braeburn.cs.iit.edu/jupyter/user/lee/notebooks/cs331/source/Lectures/priority-queues.ipynb#2.-Naive-implementation) (and the following examples) for details on how such a `key` function might be used. Your changes to the `Heap` class will need to be more extensive, however, than those found in the naive priority queue, as we are not using a built-in `sort` (which already takes a `key` function). 

Just as with the naive priority queue, the default `key` function should simply sort on the value of the elements themselves — such a default value has already been inserted for you into the `__init__` parameter list. You will, at the very least, need to update the `_heapify` and `add` methods, below, to complete this assignment.

<!-- grade_id: heap, points: 0, grade: false, solution: true -->
```python
class Heap:
    def __init__(self, key=lambda x:x):
        self.data = []
        self.key  = key

    @staticmethod
    def _parent(idx):
        return (idx-1)//2 #the formula to get the parent
        
    @staticmethod
    def _left(idx):
        return idx*2+1

    @staticmethod
    def _right(idx):
        return idx*2+2
    
    def _heapify(self, curr=0):
        left=Heap._left(curr)
        right=Heap._right(curr)
        max_idx=curr
        if left<len(self.data) and self.key(self.data[left])>self.key(self.data[max_idx]):
            max_idx=left
        if right<len(self.data) and self.key(self.data[right])> self.key(self.data[max_idx]):
            max_idx=right
        if max_idx !=curr:
            self.data[curr],self.data[max_idx]=self.data[max_idx],self.data[curr]
            curr=max_idx
            self._heapify(max_idx)
    
            
    def add(self, x):
        self.data.append(x)
        idx=len(self.data)-1#the last one in list
        pidx=Heap._parent(idx)
        while idx!=0 and self.key(self.data[pidx])<self.key(self.data[idx]):
            self.data[pidx],self.data[idx]=self.data[idx],self.data[pidx]
            idx=pidx
            pidx=Heap._parent(idx)
            
        
    def max(self):
        return self.data[0]

    def pop_max(self):
        ret = self.data[0]
        self.data[0] = self.data[len(self.data)-1]
        del self.data[len(self.data)-1]
        self._heapify()
        return ret
    
    def __bool__(self):
        return len(self.data) > 0

    def __len__(self):
        return len(self.data)

    def __repr__(self):
        return repr(self.data)
```

```python
h=Heap()
for i in range(10):
    h.add(i)
h

```

Output:
```
[9, 8, 5, 6, 7, 1, 4, 0, 3, 2]
```

<!-- grade_id: heap_test_1, points: 3, grade: true, solution: false -->
```python
from unittest import TestCase
import random

tc = TestCase()
h = Heap()

random.seed(0)
for _ in range(10):
    h.add(random.randrange(100))

tc.assertEqual(h.data, [97, 61, 65, 49, 51, 53, 62, 5, 38, 33])
```

<!-- grade_id: heap_test_2, points: 3, grade: true, solution: false -->
```python
from unittest import TestCase
import random

tc = TestCase()
h = Heap(lambda x:-x)

random.seed(0)
for _ in range(10):
    h.add(random.randrange(100))

tc.assertEqual(h.data, [5, 33, 53, 38, 49, 65, 62, 97, 51, 61])
```

<!-- grade_id: heap_test_3, points: 3, grade: true, solution: false -->
```python
from unittest import TestCase
import random

tc = TestCase()
h = Heap(lambda s:len(s))

h.add('hello')
h.add('hi')
h.add('abracadabra')
h.add('supercalifragilisticexpialidocious')
h.add('0')

tc.assertEqual(h.data,
              ['supercalifragilisticexpialidocious', 'abracadabra', 'hello', 'hi', '0'])
```

<!-- grade_id: heap_test_4, points: 3, grade: true, solution: false -->
```python
from unittest import TestCase
import random

tc = TestCase()
h = Heap()

random.seed(0)
lst = list(range(-1000, 1000))
random.shuffle(lst)

for x in lst:
    h.add(x)

for x in range(999, -1000, -1):
    tc.assertEqual(x, h.pop_max())
```

<!-- grade_id: heap_test_5, points: 3, grade: true, solution: false -->
```python
from unittest import TestCase
import random

tc = TestCase()
h = Heap(key=lambda x:abs(x))

random.seed(0)
lst = list(range(-1000, 1000, 3))
random.shuffle(lst)

for x in lst:
    h.add(x)

for x in reversed(sorted(range(-1000, 1000, 3), key=lambda x:abs(x))):
    tc.assertEqual(x, h.pop_max())
```