# Preliminaries

All assignments for this class will be fetched, worked-on, and submitted via a dynamic in-browser Python development and testing environment known as the [Jupyter Notebook](http://jupyter.org/).

This first machine problem is intended to familiarize you with the environment and assignment retrieval/submission system.

## Fetching assignments

If you're reading this, then you're likely already familiar with the procedure for fetching an assignment. To recap:

- After logging in using your @hawk.iit.edu account, you'll either click "Start Server" or "My Server" to go to the Jupyter dashboard
- From there, go to the "Assignments" tab, where you'll see any newly released assignments in the "Released assignments" area. You can fetch the files by clicking the "Fetch" button next to each assignment.
- After an assignment has been fetched, you'll see it in the "Downloaded assignments" area, where you can expand the assignment to click on individual notebooks to open them up. You can also validate and submit your work here. Alternatively, you can also browse into an assignment folder from the "Files" tab and click on notebook files to open them up there.

## In-browser development

One of the main reasons we're using the Jupyer Notebook platform is because it provides every student with a consistent, in-browser development environment. Each cell that you see in this notebook, in fact, can be used to evaluate Python code. Before proceeding, you might want to check out the [official Notebook interface documentation](http://jupyter-notebook.readthedocs.org/en/latest/notebook.html#notebook-user-interface).

The most important thing for now --- we'll be going over more details during lecture --- is that you can select a code cell (marked with a `In [XX]` indicator to its left), hit `Enter` to edit it, and `shift-Enter` to evaluate its contents.

Try this with the cells below. Note that definitions "stick" across cells, allowing us to write lengthy programs across cells in a notebook. Feel free to play around a bit.

```python
def say_hello():
    print("hello worlsdd!")
```

```python
say_hello()
```

### Autograded cells

Many of the machine problems you'll be working on will have auto-graded pieces. Typically, these involve functions you need to complete, paired with test-cases written in advance. 

Below, for instance, is a stubbed out `factorial` function for you to complete. If you've forgotten, the factorial of a positive integer *n* is simply the product of all integers from *n* down through 1. E.g., the factorial of $5 = 5 \times 4 \times 3 \times 2 \times 1 = 120$.

Go ahead and complete the implementation of `factorial` now. (You should delete the comment and the `raise` statement.)

<!-- grade_id: simple_math, points: 0, grade: false, solution: true -->
```python
def factorial(n):
    if(n==1):
        return n
    else:
        return factorial(n-1)*n
```

When you're done with your solution, you can run the tests in the cell below (but don't modify its contents!) to see how you did. Any errors will be displayed as output; if there is no output, it means you passed all the autograder tests.

<!-- grade_id: simple_math_test, points: 2, grade: true, solution: false -->
```python
import unittest
tc = unittest.TestCase()
tc.assertEqual(factorial(1), 1)
tc.assertEqual(factorial(5), 120)
tc.assertEqual(factorial(10), 3628800)
```

```python
#the code above is the code to test the factorial function
```

### Project Euler problems

Because we've yet to discuss any data structures or algorithms, these next few general programming problems are taken from [Project Euler](https://projecteuler.net/) to provide some warm-up exercises for the Python language.

#### Multiples of 3 and 5

If we list all the natural numbers below 10 that are multiples of 3 or 5, we get 3, 5, 6 and 9. The sum of these multiples is 23.

Find the sum of all the multiples of 3 or 5 below the provided argument `n`.

<!-- grade_id: multiples_of_3_and_5, points: 0, grade: false, solution: true -->
```python
def multiples_of_3_and_5(n):
    m=0
    z=0
    for i in range(n):
        if(i%3==0):
            m=i+m
            i=i+3
        elif(i%5==0):
            z=i+z
            i=i+5
    return (m+z)
```

```python
#m and z here is to add up total sum
```

<!-- grade_id: multiples_of_3_and_5_test, points: 4, grade: true, solution: false -->
```python
import unittest
tc = unittest.TestCase()
tc.assertEqual(multiples_of_3_and_5(10), 23)
tc.assertEqual(multiples_of_3_and_5(500), 57918)
tc.assertEqual(multiples_of_3_and_5(1000), 233168)
```

Output:
```
<ipython-input-4-c989b6807c69> in <module>()
      1 import unittest
----> 2 tc = unittest.TestCase()
      3 tc.assertEqual(multiples_of_3_and_5(10), 23)
      4 tc.assertEqual(multiples_of_3_and_5(500), 57918)
      5 tc.assertEqual(multiples_of_3_and_5(1000), 233168)

/Users/huangzhen/anaconda/lib/python2.7/unittest/case.pyc in __init__(self, methodName)
    187         except AttributeError:
    188             raise ValueError("no such test method in %s: %s" %
--> 189                   (self.__class__, methodName))
    190         self._testMethodDoc = testMethod.__doc__
    191         self._cleanups = []

ValueError: no such test method in <class 'unittest.case.TestCase'>: runTest
```

#### Even Fibonacci numbers

Each new term in the Fibonacci sequence is generated by adding the previous two terms. By starting with 1 and 2, the first 10 terms will be:

    1, 2, 3, 5, 8, 13, 21, 34, 55, 89, ...

By considering the terms in the Fibonacci sequence whose values do not exceed the argument `n`, find the sum of the even-valued terms.

<!-- grade_id: fib, points: 0, grade: false, solution: true -->
```python
def even_fibonacci(n):
    # YOUR CODE HERE
    pre=1;
    i=2;
    sub=0;
    result=0;
    while(i<=n):
        if(i%2==0):
            result=result+i
        sub=pre
        pre=i
        i=i+sub
    return result       
```

```python
#the pre is the previous num, the sub is to save pre
```

<!-- grade_id: fib_tests, points: 4, grade: true, solution: false -->
```python
import unittest
tc = unittest.TestCase()
tc.assertEqual(even_fibonacci(100), 44)
tc.assertEqual(even_fibonacci(10000), 3382)
tc.assertEqual(even_fibonacci(4000000), 4613732)
```

Output:
```
/usr/lib/python3.4/unittest/case.py in assertEqual(self, first, second, msg)
    795         """
    796         assertion_func = self._getAssertEqualityFunc(first, second)
--> 797         assertion_func(first, second, msg=msg)
    798 
    799     def assertNotEqual(self, first, second, msg=None):

/usr/lib/python3.4/unittest/case.py in _baseAssertEqual(self, first, second, msg)
    788             standardMsg = '%s != %s' % _common_shorten_repr(first, second)
    789             msg = self._formatMessage(msg, standardMsg)
--> 790             raise self.failureException(msg)
    791 
    792     def assertEqual(self, first, second, msg=None):

AssertionError: 128 != 44
```

## Submission

When you're ready to submit your work, make sure it's saved (you can click the disk icon in the menu bar), then go back to the "Assignments" tab and click "Submit" next to this assignment in the "Downloaded assignments" area. You can submit as many times as you like before the due date -- you'll see your submissions listed in the "Submitted assignments" area.

That's it for now!